#### Original Author: Rens Dimmendaal (TA on the course in 2017/2018)

For this problem set, we'll be using the Jupyter notebook:

![](jupyter.png)

# Week 1: Tutorial on Matplotlib and Pandas

Welcome! In this week's tutorial you will learn to create visualizations with *matplotlib* and work with data tables with *pandas*.


**Note: We assume you are able to open and work with these type Jupyter notebooks, if that's not the case please reach out to us.**

## Useful resources

Let's start off with some useful resources, keep them ready during this tutorial:

* Cheatsheets (courtesy of datacamp.com):
    * [Matplotlib](https://s3.amazonaws.com/assets.datacamp.com/blog_assets/Python_Matplotlib_Cheat_Sheet.pdf)
    * [Pandas](https://s3.amazonaws.com/assets.datacamp.com/blog_assets/PandasPythonForDataScience.pdf)
    * [Basic python](https://s3.amazonaws.com/assets.datacamp.com/blog_assets/PythonForDataScience.pdf)
* Official tutorials (these websites also contain tbhe full documentation):
    * [Matplotlib](https://matplotlib.org/users/pyplot_tutorial.html)
    * [Pandas](http://pandas.pydata.org/pandas-docs/stable/10min.html#min)

Extra: [Matplotlib gallery with sample code](https://matplotlib.org/gallery.html)
    
    
**The idea of today's assignments is that you become self-sufficient in finding solutions to your problems, these documents and google are your new friends :-)**


## Assumptions

We assume you're running python 3 and have the ability to install packages with pip or conda where needed.

Packages you will use today:

* Matplotlib (for making plots)
* Pandas (for working with tables)
* Jupyter (for the interactive notebook)

You can load the matplotlib and pandas capabilities into your jupyter notebook as follows:

```python
import matplotlib.pyplot as plt
import pandas as pd
import numpy as np # a math library for some tests don't worry about it
```

You can setup jupyter to immideately show plots within your notebook as follows:

```python
%matplotlib inline
# % matplotlib notebook # Or alternatively a notebook version
```

Check if everything works by running the code below.
You should see a plot with three red crosses

```python
"""
TEST CODE -- CHECK MANUALLY
"""
tf = pd.DataFrame({'x':[1,2,3],'y':[10,20,10]})
plt.scatter(tf['x'],tf['y'],color='red',marker='+',s=300)
plt.show()
```

## Part (A) Matplotlib on Datacamp.com

At the lectures on Tuesday you have heard about the basics of matplotlib. We warmly recommend you to complete all exercises from the notebook used for the lecture. 

Matplotlib is also a part of the “Intermediate Python for Data Science” and “Introduction to Data Visualization with Python” modules on DataCamp. If you would like more practice with matplotlib, those would be a good starting point.

* [Matplotlib 1](https://www.datacamp.com/courses/intermediate-python-for-data-science)
* [Matplotlib 2](https://www.datacamp.com/courses/introduction-to-data-visualization-with-python)

## Part (B) Pandas

Now we will learn to work with pandas, a package to work wiht tables in python.
We will provide less 'hand holding' than Datacamp, so be sure to keep the cheat sheets ready!

### Load data


Let's start with loading some data from the CBS - Statistics Netherlands, which you can download from Canvas.

<!-- grade_id: load_csv, points: 0, grade: false, solution: true -->
```python
"""
ASSIGNMENT: Load the csv file as a pandas DataFrame named df.
HINT: Make sure you use the correct separator!
"""
df = None
### BEGIN SOLUTION
df = pd.read_csv('cbs_data.csv')
### END SOLUTION
```

<!-- grade_id: check_load_csv, points: 1, grade: true, solution: false -->
```python
"""
CHECK
"""
assert df.columns.tolist() == ['province',
                                 'year',
                                 'total_population',
                                 'nr_women',
                                 'nr_younger_than_5',
                                 'nr_between_5_and_10',
                                 'nr_between_10_and_15',
                                 'nr_between_15_and_20',
                                 'nr_between_20_and_25',
                                 'nr_between_25_and_45',
                                 'nr_between_45_and_65',
                                 'nr_between_65_and_80',
                                 'nr_above_80',
                                 'nr_in_urban_region',
                                 'nr_buildings',
                                 'nr_buildings_per_km2',
                                 'nr_personal_vehicles',
                                 'km_total_road',
                                 'km_city_road',
                                 'km_provincial_road',
                                 'km_national_road']

assert len(df) == 180
```

### Inspecting data

You can inspect the data by calling some summary functions.

Try the following:
    
```
k=5
df.head(k)
df.tail(k)
df.describe()
```

Describe what each function does:

```python
"""
ASSIGNMENT: 
    1. Try out the functions mentioned above in this cell.
    2. Describe what they do in the cell below
"""

### YOUR CODE HERE (just for trial and error)
```

It's often useful to have some lists directly available. 
Extract them now and store them in separate variables.
Create separate lists for:
1. All provinces (unique)
2. All years (unique)
3. All variables (all columns, but not provinces and years)

```python
df.columns
```

<!-- grade_id: get_vars, points: 0, grade: false, solution: true -->
```python
"""
ASSIGNMENT: Store the name of all provinces, years, and variables
"""

### BEGIN SOLUTION
provinces = df.province.unique().tolist()
years = df.year.unique().tolist()
columns = df.columns.tolist()
variables = columns[2:]
### END SOLUTION
```

```python
years == [x for x in range(2003,2018)]
```

<!-- grade_id: check_get_vars, points: 1, grade: true, solution: false -->
```python
"""
CHECK
"""

assert provinces == ['Groningen',
                    'Friesland',
                    'Drenthe',
                    'Overijssel',
                    'Flevoland',
                    'Gelderland',
                    'Utrecht',
                    'Noord-Holland',
                    'Zuid-Holland',
                    'Zeeland',
                    'Noord-Brabant',
                    'Limburg']

assert years == [x for x in range(2003,2018)]

assert variables == ['total_population',
                     'nr_women',
                     'nr_younger_than_5',
                     'nr_between_5_and_10',
                     'nr_between_10_and_15',
                     'nr_between_15_and_20',
                     'nr_between_20_and_25',
                     'nr_between_25_and_45',
                     'nr_between_45_and_65',
                     'nr_between_65_and_80',
                     'nr_above_80',
                     'nr_in_urban_region',
                     'nr_buildings',
                     'nr_buildings_per_km2',
                     'nr_personal_vehicles',
                     'km_total_road',
                     'km_city_road',
                     'km_provincial_road',
                     'km_national_road']





```

### Selecting rows

You'll often only be interested in a subset of the data. 
Let's write a couple of functions for some specific queries

<!-- grade_id: extract_province, points: 0, grade: false, solution: true -->
```python
def extract_province(df,province):
    ### BEGIN SOLUTION
    result = df.loc[df['province'] == province]
    ### END SOLUTION
    return result
```

<!-- grade_id: check_extract_province, points: 1, grade: true, solution: true -->
```python
assert df.iloc[0:15].equals(extract_province(df,provinces[0]))

### BEGIN HIDDEN TESTS
for province in ['Limburg','Noord-Holland','Zuid-Holland','Friesland']:
    assert df.loc[df['province'] == province].equals(extract_province(df,province))
### END HIDDEN TESTS
```

<!-- grade_id: extract_variables, points: 0, grade: false, solution: true -->
```python
"""
ASSIGNMENT: Create a function to extract BOTH:
    1. Variables of interest
    2. Province & Year information
"""
def extract_variables(df,variables):
    ### BEGIN SOLUTION
    result = df[['province', 'year'] + variables]
    ### END SOLUTION
    return result
```

<!-- grade_id: check_extract_variables, points: 1, grade: true, solution: false -->
```python
"""
CHECK: extract_variables()
"""
test_cols = ['total_population','km_provincial_road']
test = extract_variables(df,test_cols)
assert test.columns.tolist() == ['province','year'] + test_cols
assert len(test) == len(df)
```

### Grouping rows

We can use the groupby functionality to summarize our datasets by groups.
Use this functionality to complete the following two functions:

<!-- grade_id: year_total, points: 0, grade: false, solution: true -->
```python
"""
Assignment: Create a function which shows the total (sum) for a set of variables per year
"""

def year_total(df,variables):
    ### BEGIN SOLUTION
    result = extract_variables(df,variables)
    result = result.groupby('year').sum()
    ### END SOLUTION
    return result
    
```

```python
year_total(df,test_cols).values
```

<!-- grade_id: check_year_total, points: 1, grade: true, solution: false -->
```python
automatic = year_total(df,test_cols).values
manual = np.array([[16192572,     7858],
               [16258032,     7800],
               [16305526,     7744],
               [16334210,     7746],
               [16357992,     7898],
               [16405399,     7850],
               [16485787,     7835],
               [16574989,     7861],
               [16655799,     7863],
               [16730348,     7803],
               [16779575,     7778],
               [16829289,     7750],
               [16900726,     7738],
               [16979120,     7758],
               [17081507,     7756]])

assert np.sum(automatic == manual) == manual.shape[0] * manual.shape[1]
```

### Create new columns

You'll often want to create new columns such as ratio's, growth percentages, or other transformations.

You can simply pass a list as a new column as follows:

```
df['NEW_COLUMN_NAME'] = [x for x in range(len(df))]
```

Pandas series (one column is called a series) also understand list operations. For example:

```
df['Percentage_urban_population'] = df['Urban_population'] / df['Total_population']
```

Finally you can use a for loop to update a pandas data frame row by row. This is relatively slow, so avoid it when you can:

```
for idx, row in df.iterrows():
    df.loc[idx,'NEW_COL'] = 'foo_' + str(idx)
```


```python
"""
ASSIGNMENT: Just try and create some new variables in which you are interested
Inspiration: ratio's, log transformations, GROWTH PERCENTAGES
"""
```

### Summary 

You have now learned the basics of working with table data in pandas.
You have also learend how to lookup operations you are not familiar with in the documentation.

## Part (C) Data exploration: combine data manipulation with visualization

Now that you have learned how to plot, and how to load and manipulate data it is time to combine the two.
Explore the data and create an *effective* and *engaging* visualization that leads you to learn something new from the data.

Inspiration:

* Create a panel for each province and show the population split per age group. Which province is the greyest?
* Show how age groups change over time per province. Which provinces are 'aging' the fastest?
* Show the growth rates of each age group over time per province. Any strange results? What might cause them?

```python
# EXAMPLE
age_groups = ['nr_younger_than_5',
             'nr_between_5_and_10',
             'nr_between_10_and_15',
             'nr_between_15_and_20',
             'nr_between_20_and_25',
             'nr_between_25_and_45',
             'nr_between_45_and_65',
             'nr_between_65_and_80',
             'nr_above_80']

df['percentage_over_65'] = (df['nr_between_65_and_80'] + df['nr_above_80']) / df['total_population']


```

```python
columns = columns + ['percentage_over_65']
```

```python
col = 'percentage_over_65'

fig, ax = plt.subplots()

for province in provinces:
    temp = df.loc[df.province == province]
    ax.plot(temp['year'],temp[col],label=province)
legend = ax.legend(loc='upper left', shadow=True)
frame = legend.get_frame()


```

```python
col = 'nr_buildings_per_km2'

fig, ax = plt.subplots()

sorted_provinces = df.loc[df.year == 2017].sort_values(by=col,ascending=False)['province'].tolist()

for province in sorted_provinces:
    temp = df.loc[df.province == province].sort_values(by='province',)
    ax.plot(temp['year'],temp[col],label=province)
legend = ax.legend(loc='upper left', shadow=True)
frame = legend.get_frame()


```

```python
fig, ax = plt.subplots()

for province in sorted_provinces:
    temp = df.loc[df.province == province].sort_values(by='province',)
    ax.scatter(temp['nr_buildings_per_km2'],temp['percentage_over_65'],label=province)
legend = ax.legend(loc='upper left', shadow=True,bbox_to_anchor=(1,1))
frame = legend.get_frame()
```

```python

```